---
title: 'Linear Regration'
description: 'Linear Regration'
---

## Linear regression

```yaml
type: VideoExercise
key: 072f1a0a17
xp: 50
```

`@projector_key`
582144798ebcb5723383e7452eba64f5

---

## Calculate the correlation

```yaml
type: MultipleChoiceExercise
key: 4ddd1526c4
xp: 50
```

Get the data file from the following link:
https://tinyurl.com/weightdata

- Open a new jupyter notebook and load the data into a dataframe
- Plot a scatter plot of the weight vs the height
- Calculate the correlation between them

`@possible_answers`
- [0.924]
- 0.927
- 0.973
- 0.945

`@hint`
<!-- Examples of good hints: https://instructor-support.datacamp.com/en/articles/2379164-hints-best-practices. -->
- This is an example hint.
- This is an example hint.

`@pre_exercise_code`
```{python}

```

`@sct`
```{python}
# Check https://instructor-support.datacamp.com/en/articles/2375523-course-multiple-choice-with-console-exercises on how to write feedback messages for this exercise.
```

---

## Find the analitical fit

```yaml
type: PureMultipleChoiceExercise
key: 483dc017d0
xp: 50
```

Use the same data as before
Use the following formula to calculate the linear fit
![https://assets.datacamp.com/production/repositories/5585/datasets/b71d3b3278f46116384c60a0ad53d84d431bc783/1_O2xXpYS7A3VlLXqT6SNMLg.png]()
- What is the value of b1   

`@hint`


`@possible_answers`
- [7.17]
- 7.64
- 8.9
- 9.99

`@feedback`
<!-- Examples of good feedback messages: https://instructor-support.datacamp.com/en/articles/2299773-exercise-success-messages.  -->
- Perfect!
- Wrong answer
- Wrong answer
- Wrong answer

---

## Calculate RMSE

```yaml
type: PureMultipleChoiceExercise
key: 66c72d4f7f
xp: 50
```

Use the following formula to calculate the Root Mean squared Error RMSE
![https://assets.datacamp.com/production/repositories/5585/datasets/6d1423f9241cd62e1ed6aa177405c74d3c99d572/1_SGBsn7WytmYYbuTgDatIpw.gif]()

`@hint`
<!-- Examples of good hints: https://instructor-support.datacamp.com/en/articles/2379164-hints-best-practices. -->
- This is an example hint.
- This is an example hint.

`@possible_answers`
- 845.7
- 836.5
- [864.7]
- 832.3
- 824.8

`@feedback`
<!-- Examples of good feedback messages: https://instructor-support.datacamp.com/en/articles/2299773-exercise-success-messages.  -->
- Wrong answer
- Wrong answer
- Perfect!
- Wrong answer
- Wrong answer
