---
title: 'Linear Regration'
description: 'Linear Regration'
---

## Linear regression

```yaml
type: VideoExercise
key: 072f1a0a17
xp: 50
```

`@projector_key`
582144798ebcb5723383e7452eba64f5

---

## Calculate the correlation

```yaml
type: MultipleChoiceExercise
key: 4ddd1526c4
xp: 50
```

Get the data file from the following link:
https://tinyurl.com/weightdata

- Open a new jupyter notebook and load the data into a dataframe
- Plot a scatter plot of the weight vs the height
- Calculate the correlation between them

`@possible_answers`
- [0.924]
- 0.927
- 0.973
- 0.945

`@hint`
<!-- Examples of good hints: https://instructor-support.datacamp.com/en/articles/2379164-hints-best-practices. -->
- This is an example hint.
- This is an example hint.

`@pre_exercise_code`
```{python}

```

`@sct`
```{python}
# Check https://instructor-support.datacamp.com/en/articles/2375523-course-multiple-choice-with-console-exercises on how to write feedback messages for this exercise.
```
